```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

# Homework 02: Arrays and Tables

Please complete this notebook by filling in the cells provided. Before you begin, execute the following cell to load the provided modules. Each time you start your server, you will need to execute this cell again to load the modules.

```python
# Don't change this cell; just run it. 
import numpy as np
from datascience import *
```

**Helpful Resource:**
- [Python Reference](http://data8.org/sp22/python-reference.html): Cheat sheet of helpful array & table methods used in this course.

**Recommended Reading:**
- [What is Data Science](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)
- [Causality and Experiments](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html) 
- [Programming in Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

<div class="alert alert-warning" role="alert"  style="color: black;">
    
### Appropriate collaboration and AI usage

Consult with your teacher and course syllabus for information and policies regarding appropriate collaboration with other students, appropriate use of AI tools, and submission of late work.</div>

<div class="alert alert-danger" role="alert"  style="color: black;">
    
### This homework has hidden tests in it!

Additional tests will be run once your homework is submitted for grading. While you may pass all the tests you have access to before submission, you may not earn full credit if you do not pass the hidden tests as well.

Many of the tests you have access to before submitting only test to ensure you have given an answer that is formatted correctly and/or you have given an answer that *could* make sense in context. For example, a test you have access to while completing the assignment may check that you selected a valid choice for a multiple choice problem (1, 2, or 3) or that your answer is an integer between 0 and 50 if asked to count a subset of states in the United States. The tests that are run after submission will evaluate your work for accuracy. **Do not assume that just because all your tests pass before submission means that your answers are correct!**
</div>

## 1. Creating Arrays

### Question 1.1

Make an array called `planets` containing the following eight strings in this order: "Mercury", "Venus", "Earth", "Mars", "Jupiter", "Saturn", "Uranus", "and Neptune".

```python
planets = ...
planets
```

```python
grader.check("q1_1")
```

### Question 1.2

Make an **array** called `just_numbers` containing the following numbers (in the given order):

1. $3^{\sin(7)}$
2. $81^{0.25}$
3. $\tan(45)$
4. $-12$
   

**Hint:** `sin` and `tan` are functions in the `math` module. Importing modules is covered in a previous lab! You should import the module in your solution cell below.

**Note:** Python lists are different/behave differently than numpy arrays. In this course, we use `numpy` arrays, so please make an **array**, not a python list.

```python
...
just_numbers = ...
just_numbers
```

```python
grader.check("q1_2")
```

### `join`
All strings in Python have a built-in method called `join`. The `join` method takes only one argument as an input, a collection (like an array) of strings. The `join` method returns a single string as an output.

The output string of `a_string.join(an_array)` is a single string that's the [concatenation](https://en.wikipedia.org/wiki/Concatenation) ("putting together") of all the strings in `an_array`, with `a_string` inserted in between each string from the array. Run the cell below to see it in action:

```python
# Just run this cell
example_array = make_array('first', 'second', 'third')
' and then '.join(example_array)
```

### Question 1.3

Use the array `planets` and the method `join` to make two strings:

1. "Mercury Venus Earth Mars Jupiter Saturn Uranus and Neptune" (call this one `without_commas`)
2. "Mercury, Venus, Earth, Mars, Jupiter, Saturn, Uranus, and Neptune" (call this one `with_commas`)

**Hint:** If you're still not sure what `join` does, first try just calling, for example, `"SPACE".join(planets)` in a new code cell below.

```python
without_commas = ...
with_commas = ...
print(without_commas)
print(with_commas)
```

```python
grader.check("q1_3")
```

## 2. Indexing Arrays

These exercises give you practice accessing individual elements of arrays. In Python (and in many programming languages), elements are accessed by *index*, so the first element is the element at index 0.

<div class="alert alert-warning" role="alert"  style="color: black;">

**Note:** If you have previous coding experience, you may be familiar with bracket notation. **DO NOT** use bracket notation when indexing (i.e. `arr[0]`), as this can yield different data type outputs than what we will be expecting. This can cause you to fail an autograder test.

</div>

Be sure to refer to the [Python Reference](http://data8.org/sp22/python-reference.html) on the website if you feel stuck!

### Question 2.1

Create an array called `some_numbers` that includes the numbers 25, -7, -15, 42, 88, in that order. Set `third_element` to the third element of `some_numbers` using the `item` method.

```python
some_numbers = ...
third_element = ...
third_element
```

```python
grader.check("q2_1")
```

### Question 2.2

The next cell creates a table that displays some information about the elements of `some_numbers` and their order. Run the cell to see the partially-completed table, then fill in the missing information (the cells that say "Ellipsis") by assigning `blank_a`, `blank_b`, `blank_c`, and `blank_d` to the correct elements in the table.

**Hint:** Replace the ... with strings or numbers to make the Table consistent with the array from the previous problem.

```python
blank_a = ...
blank_b = ...
blank_c = ...
blank_d = ...

Table().with_columns(
    'English name for position', make_array('first', blank_a, 'third', 'fourth', 'fifth'),
    'Index', make_array(blank_b, 1, 2, blank_c, 4),
    'Element', make_array(25, -7, blank_d, 42, 88))
```

```python
grader.check("q2_2")
```

<!-- BEGIN QUESTION -->

### Question 2.3

A student claims that "the fifth element of an array is at index 5." Explain why this statement is incorrect, and describe a common mistake that might lead someone to think this way.

<div class="alert alert-info" role="alert" style="color: black;">

Write your explanation here, replacing this text.

</div>

<!-- END QUESTION -->

## 3. Basic Array Arithmetic

### Question 3.1

Divide the numbers `36`, `-4,284`, `1,296,000,000`, and `720` by `12`. Assign each variable below such that `first_quotient` is assigned to the result of $36 \div 12$, `second_quotient` is assigned to the result of $-4,284 \div 12$, and so on.

For this question, **don't** use arrays.

```python
first_quotient = ...
second_quotient = ...
third_quotient = ...
fourth_quotient = ...
print(first_quotient, second_quotient, third_quotient, fourth_quotient)
```

```python
grader.check("q3_1")
```

### Question 3.2

Now, do the same calculation, but using an array called `numbers` that contains the 4 original numbers (`36`, `-4,284`, `1,296,000,000`, and `720`) and only a single division (`/`) operator. Store the 4 results in an array named `quotients`.

```python
numbers = ...
quotients = ...
quotients
```

```python
grader.check("q3_2")
```

### Question 3.3

Oops, we made a typo! Instead of dividing by 12, we wanted to **multiply** each number by 12. Compute the correct products in the cell below using array arithmetic and store them in an array called `products`. The purpose of this question is to help you realize that it's much easier to perform (and in this case, change) arithmetic on an array all at once, rather than having to do so once for each element in the array. Notice that your job is going to be really easy if you previously defined an array containing the 4 numbers.

```python
products = ...
products
```

```python
grader.check("q3_3")
```

## 4. Spotify Top Songs

The file `spotify_top_songs.csv` contains data about some of Spotify's most-streamed songs of all time. Run the next cell to load the data.

```python
# Just run this cell
songs_info = Table.read_table('spotify_top_songs.csv')
songs_info
```

### Question 4.1

The cell below selects the column labeled `'streams_millions'` from the `songs_info` Table and assigns it to the array called `song_streams`. This column contains the total streams (in millions) for each song. **Reminder:** Calling `.column(...)` on a table returns an array of the column specified, in this case the `streams_millions` column of the `songs_info` table.

Use the `max` function to assign the largest stream count found in the Table to `highest_streams`.

```python
song_streams = songs_info.column('streams_millions')
```

```python
highest_streams = ...
highest_streams
```

```python
grader.check("q4_1")
```

<!-- BEGIN QUESTION -->

### Question 4.2

There are several other ways you could have determined the maximum stream count in the `songs_info` Table without using the `max` function. For example, there's one way that utilizes both the `len(...)` and `.item(...)` functions (as well as a few other commands),that don't require you to know the number of elements in `song_streams` to be able to determine the value.

First, explain in words how you could use Table methods (for example, `.sort`), in combination with other array functions and methods to compute the maximum number of streams in the Table. 

Then, implement your solution using Python code in the provided cell. Your code should be able to work no matter how big or small the `song_streams` array might be. Compare your results to the answer in the previous question to help check that your code is working properly.

<div class="alert alert-info" role="alert" style="color: black;">

Write your explanation here, replacing this text.

</div>

```python
# Write you code in this cell
...
```

<!-- END QUESTION -->

### Question 4.3

The average song in our dataset has about 2,150 million streams. What percentage of this average does each song represent? For example, a song with 4,300 million streams would represent 200% of the average. Compute this percentage for each song and round to two decimal places. Store the values in the array named `percent_of_average`.

```python
percent_of_average = ...
percent_of_average
```

```python
grader.check("q4_3")
```

<!-- BEGIN QUESTION -->

### Question 4.4

Look at the `danceability` and `energy` columns in the `songs_info` table. Both are measured on a scale from 0 to 1.

A music blogger claims: "Songs with high danceability always have high energy." 

Based on what you can observe in the data, do you agree or disagree? Use specific examples from the table to support your answer.

```python
# Write some code here to explore the data to help formulate a written response
...
```

<div class="alert alert-info" role="alert" style="color: black;">

Write your explanation here, replacing this text.

</div>

<!-- END QUESTION -->

## 5. Netflix Subscriber Growth

The following table shows Netflix's global subscriber count (in millions) from 2010 to 2024. Run the cell below to load the data.

```python
# Just run this cell
netflix = Table.read_table('netflix_subscribers.csv')
netflix
```

### `numpy` functions

In upcoming questions, you will apply some built-in `numpy` functions to this array. `numpy` is a Python module that is often used in Data Science. This homework will use two different `numpy` functions, `np.diff` and `np.cumsum`. **Note:** While used on this homework, these specific functions are not used heavily throughout the course. This set of questions is meant to get you quickly learning and using new functions in your programming, which is something data scientists need to do all the time!

#### `np.diff`
The difference function `np.diff` subtracts each element in an array from the element after it within the array. As a result, the length of the array `np.diff` returns will always be one less than the length of the input array.

The diagram below illustrates how this function operates.
<img src="array_diff.png" style="width: 600px;"/>

#### `np.cumsum`

The **cumulative** sum function `np.cumsum` outputs an array of partial sums. For example, the third element in the output array corresponds to the sum of the first, second, and third elements.

The diagram below illustrates how this function operates.

<img src="array_cumsum.png" style="width: 700px;"/>

### Question 5.1

The first line below assigns `subscribers` to an array of 15 subscriber counts from consecutive years of Netflix data; specifically, 2010-2024. Assign the names `least`, `greatest`, and `average` so that the `print` statement is correct.

```python
subscribers = netflix.column('subscribers_millions')
```

```python
least = ...
greatest = ...
average = ...
print('The least subscribers in a year was', least, 'million.')
print('The greatest subscribers in a year was', greatest, 'million.')
print('The average subscribers per year was', average, 'million.')
```

```python
grader.check("q5_1")
```

### Question 5.2

Assign `biggest_increase` to the biggest increase in subscribers between two consecutive years. For example, the second year in the data set has 23.53 million subscribers and the third has 27.15 million, so the increase was $27.15 - 23.53 = 3.62$ million subscribers.

**Hint:**

- You'll need an array arithmetic function [mentioned in the textbook](https://inferentialthinking.com/chapters/05/1/Arrays.html#functions-on-arrays).

```python
biggest_increase = ...
biggest_increase
```

```python
grader.check("q5_2")
```

### Question 5.3

You determined in an earlier question that the average subscriber count during this time period was around 75 million. Therefore, you might *expect* the **total cumulative** subscriber-years to be `75 * n` million after `n` years during this 15 year period.

1. Use the `np.arange` function to create an array of values where each element represents the **total cumulative** subscriber-years (in millions) one would expect Netflix to have accumulated by the end of each year, for 15 years. Use the assumption that they have 75 million subscribers each year. For example, the first element would be 75, the second element 150, the third 225, and so on. Assign this array of values to `expected_subscribers`. You can read about `np.arange` as needed in this [textbook section](https://inferentialthinking.com/chapters/05/2/Ranges.html).

2. Then, determine how the actual total cumulative subscriber-years over the 15 year period compared to what we expected to happen under our assumption. Set `difference_from_expected` to an array of 15 elements, where each element is the **absolute difference** between the actual cumulative subscribers and expected cumulative subscribers in that particular year.

For example, since the first three subscriber counts are 20.01, 23.53, and 27.15 million. The total cumulative for 3 years is 20.01 + 23.53 + 27.15 = 70.69. The expected for 3 years is 75 * 3 = 225. Therefore, `difference_from_expected.item(2)` should be $|70.69 - 225| = 154.31$.

**Hint 1**: You'll need to use the `np.cumsum` function to create the array of actual cumulative subscribers.

**Hint 2**: When using the absolute value function, use `np.abs` instead of `abs`, since `numpy` is already loaded in this notebook, and that's what the autograder is going to use.

```python
expected_subscribers = ...
actual_cumulative = ...
difference_from_expected = ...
difference_from_expected
```

```python
grader.check("q5_3")
```

### Question 5.4

Let's imagine your guess for the next year's subscriber count was always just the subscriber count of the previous year. If you always guessed the previous year's count, how big would your error in guessing be, on average? Compute the value using some `numpy` operations and assign your answer to `average_error`.

For example, the first three subscriber counts are 20.01, 23.53, and 27.15. After the first year you guess the next year's count will be 20.01 million, but it is 23.53 million resulting in an error of $|20.01 - 23.53|$. For the next year you now guess there will be 23.53 million subscribers, but there are 27.15 million, so the error is $|23.53 - 27.15|$. Therefore, the average error for just the second and third years would be $$\frac{|20.01-23.53| + |23.53-27.15|}{2} = 3.52$$.

```python
average_error = ...
average_error
```

```python
grader.check("q5_4")
```

## 6. Tables

### Question 6.1

Suppose you have 5 science fiction books, 3 mystery books, and 4 biography books. (Perhaps you're organizing your bookshelf.) Create a table that contains this information. It should have two columns: `genre` and `count`. Assign the new table to the variable `books`.

**Note:** Use lower-case for the genre names, like `"scifi"`.

Your table should look like:

| genre    | count |
|----------|-------|
| scifi    | 5     |
| mystery  | 3     |
| biography| 4     |

```python
books = ...
    ...
    ...
books
```

```python
grader.check("q6_1")
```

### Question 6.2

The file `bookstore_inventory.csv` contains information about the inventory at a bookstore. Each row represents a different book title. Load it as a table named `inventory` using the `Table.read_table()` function. `Table.read_table(...)` takes one argument (data file name in string format) and returns a table.

```python
inventory = ...
inventory
```

```python
grader.check("q6_2")
```

### Question 6.3

The file `bookstore_sales.csv` contains the number of copies sold of each book last month. It has an extra column called "price_per_copy" that's the price for each book. The rows are in the same order as the `inventory` table.

Load these data into a table called `sales`.

```python
sales = ...
sales
```

```python
grader.check("q6_3")
```

### Question 6.4

How many books did the store sell in total last month?

```python
total_books_sold = ...
total_books_sold
```

```python
grader.check("q6_4")
```

### Question 6.5

What was the store's total revenue (the total price of all books sold) last month?

**Hint:** If you're stuck, think first about how you would compute the total revenue from just the Harry Potter sales.

```python
total_revenue = ...
total_revenue
```

```python
grader.check("q6_5")
```

### Question 6.6

Make a new table called `remaining_inventory`. It should have the same rows and columns as `inventory`, except that the number of copies sold should be subtracted from each book's **original** stock, so that the `stock` column is **updated to be** the number of copies remaining after last month's sales.

<div class="alert alert-warning" role="alert"  style="color: black;">

**HINT:** You can take data from an existing table, like `inventory`, and use the `.with_column(...)` method to overwrite an existing column of data. All you need to do is use `.with_column(...)` as if you were adding a new column, but just specify and existing column label.

For example:

```python
remaining_inventory = inventory.with_column('stock', make_array(0,0,0,0,0,0,0,0))

```


takes the data from the `inventory` Table, sets every value in the 'stock' column to 0, and saves the result to `remaining_inventory`.

</div>

```python
remaining_inventory = ...
remaining_inventory
```

```python
grader.check("q6_6")
```

# Submitting your work
You're done with this assignment! Assignments should be turned in using the following best practices:
1. Save your notebook.
2. Restart the kernel and run all cells up to this one.
3. Run the cell below with the code `grader.export(...)`. This will re-run all the tests. Make sure they are passing as you expect them to.
4. Download the file named `hw02_<date-time-stamp>.zip`, found in the explorer pane on the left side of the screen. **Note**: Clicking on the link in this notebook may result in an error, it's best to download from the file explorer panel.
5. Upload `hw02_<date-time-stamp>.zip` to the corresponding assignment on Canvas.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(force_save=True, pdf=False)
```

Tests for q1_1 (points: None)
```
>>> isinstance(planets, np.ndarray)
True
```
failure_message: ❌ You did not create an array. Perhaps you created a list or another type of collection?
success_message: ✅ You successfully created an array.
```
>>> not any([',' in text for text in planets])
True
```
failure_message: ❌ You included commas inside some of the strings. Commas are used to separate individual strings, but are not between the quotes.
```
>>> 'and ' in planets.item(7)
True
```
failure_message: ❌ Your last string does not contain the substring `and `
success_message: ✅ The last string contains the substring `and `


Tests for q1_2 (points: None)
```
>>> isinstance(just_numbers, np.ndarray)
True
```
failure_message: ❌ You did not create an array. Perhaps you created a list or another type of collection?
success_message: ✅ You successfully created an array.
```
>>> len(just_numbers) == 4
True
```
failure_message: ❌ Your array should have four elements.
success_message: ✅ Your array has four elements.


Tests for q1_3 (points: None)
```
>>> isinstance(without_commas, str)
True
```
failure_message: ❌ without_commas should be a string
success_message: ✅ without_commas is a string


Tests for q2_1 (points: None)
```
>>> len(some_numbers) == 5
True
```
failure_message: ❌ Your array should have 5 elements
success_message: ✅ You created an array with 5 elements


Tests for q2_2 (points: None)
```
>>> isinstance(blank_a, str)
True
```
failure_message: ❌ blank_a should be a string
success_message: ✅ blank_a is a string


Tests for q3_1 (points: None)
```
>>> isinstance(first_quotient, (int, float))
True
```
failure_message: ❌ first_quotient should be a number
success_message: ✅ first_quotient is a number


Tests for q3_2 (points: None)
```
>>> isinstance(quotients, np.ndarray)
True
```
failure_message: ❌ quotients should be an array
success_message: ✅ quotients is an array


Tests for q3_3 (points: None)
```
>>> isinstance(products, np.ndarray) and len(products) == 4
True
```
failure_message: ❌ products should be an array with 4 elements
success_message: ✅ products is an array with 4 elements


Tests for q4_1 (points: None)
```
>>> isinstance(highest_streams, (int, float, np.integer, np.floating))
True
```
failure_message: ❌ highest_streams should be a number
success_message: ✅ highest_streams is a number


Tests for q4_3 (points: None)
```
>>> isinstance(percent_of_average, np.ndarray)
True
```
failure_message: ❌ percent_of_average should be an array
success_message: ✅ percent_of_average is an array


Tests for q5_1 (points: None)
```
>>> isinstance(least, (int, float, np.integer, np.floating))
True
```
failure_message: ❌ least should be a number
success_message: ✅ least is a number


Tests for q5_2 (points: None)
```
>>> isinstance(biggest_increase, (int, float, np.integer, np.floating))
True
```
failure_message: ❌ biggest_increase should be a number
success_message: ✅ biggest_increase is a number


Tests for q5_3 (points: None)
```
>>> isinstance(difference_from_expected, np.ndarray) and len(difference_from_expected) == 15
True
```
failure_message: ❌ difference_from_expected should be an array with 15 elements
success_message: ✅ difference_from_expected is an array with 15 elements


Tests for q5_4 (points: None)
```
>>> isinstance(average_error, (int, float, np.integer, np.floating))
True
```
failure_message: ❌ average_error should be a number
success_message: ✅ average_error is a number


Tests for q6_1 (points: None)
```
>>> isinstance(books, Table)
True
```
failure_message: ❌ books should be a Table
success_message: ✅ books is a Table


Tests for q6_2 (points: None)
```
>>> isinstance(inventory, Table)
True
```
failure_message: ❌ inventory should be a Table
success_message: ✅ inventory is a Table


Tests for q6_3 (points: None)
```
>>> isinstance(sales, Table)
True
```
failure_message: ❌ sales should be a Table
success_message: ✅ sales is a Table


Tests for q6_4 (points: None)
```
>>> isinstance(total_books_sold, (int, float, np.integer, np.floating))
True
```
failure_message: ❌ total_books_sold should be a number
success_message: ✅ total_books_sold is a number


Tests for q6_5 (points: None)
```
>>> isinstance(total_revenue, (int, float, np.integer, np.floating))
True
```
failure_message: ❌ total_revenue should be a number
success_message: ✅ total_revenue is a number


Tests for q6_6 (points: None)
```
>>> remaining_inventory.num_columns == 4
True
```
failure_message: ❌ Your table should have 4 columns
success_message: ✅ Your table has the correct number of columns
```
>>> remaining_inventory.column('stock').item(0) != 45
True
```
failure_message: ❌ You forgot to subtract off the sales
success_message: ✅ You remembered to subtract off the sales
