```python
# Initialize Otter
import otter
grader = otter.Notebook("hw03.ipynb")
```

# Homework 03: Table Manipulation and Visualization

Please complete this notebook by filling in the cells provided.

```python
# Don't change this cell; just run it.
import numpy as np
from datascience import *

import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

**Helpful Resource:**
- [Python Reference](http://data8.org/sp22/python-reference.html): Cheat sheet of helpful array & table methods used in this course.

**Recommended Reading:**

* [Visualization](https://inferentialthinking.com/chapters/07/Visualization.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

<div class="alert alert-warning" role="alert"  style="color: black;">
    
### Appropriate collaboration and AI usage

Consult with your teacher and course syllabus for information and policies regarding appropriate collaboration with other students, appropriate use of AI tools, and submission of late work.</div>

<div class="alert alert-danger" role="alert"  style="color: black;">
    
### This homework has hidden tests in it!

Additional tests will be run once your homework is submitted for grading. While you may pass all the tests you have access to before submission, you may not earn full credit if you do not pass the hidden tests as well.

Many of the tests you have access to before submitting only test to ensure you have given an answer that is formatted correctly and/or you have given an answer that *could* make sense in context. For example, a test you have access to while completing the assignment may check that you selected a valid choice for a multiple choice problem (1, 2, or 3) or that your answer is an integer between 0 and 50 if asked to count a subset of states in the United States. The tests that are run after submission will evaluate your work for accuracy. **Do not assume that just because all your tests pass before submission means that your answers are correct!**
</div>

## 1. Inflation

The Consumer Price Index (CPI) is a measure of the average change over time in the prices paid by consumers for a basket of goods and services. The Bureau of Labor Statistics publishes CPI data monthly.

The data file `inflation.csv` contains two key measures:
- **CPI**: The overall Consumer Price Index including all items
- **CPI-Core**: The "core" CPI which excludes volatile food and energy prices

The difference between CPI and CPI-Core can tell us about how much food and energy prices are affecting overall inflation.

#### Question 1.1.

The data are provided in a CSV file called `inflation.csv`. Load that file into a table called `inflation`.

```python
inflation = ...
inflation
```

```python
grader.check("q1_1")
```

#### Question 1.2.

Sort the data in descending order by CPI, naming the sorted table `by_cpi`. Create another table called `by_cpi_core` that's sorted in descending order by CPI-Core instead.

```python
by_cpi = ...
by_cpi_core = ...
by_cpi
```

```python
grader.check("q1_2")
```

#### Question 1.3.

Use `take` to make a table that only has the `Date` and `CPI` columns which contains only the data for the 10 months when CPI was the greatest. Call the new Table `greatest_cpi` and sort it in descending order based on the `CPI` column. Recall that each row of `inflation` represents a month.

**Hint:** You will need to remember / look up how to keep only certain columns and rows from a Table.

```python
greatest_cpi = ...
greatest_cpi
```

```python
grader.check("q1_3")
```

#### Question 1.4.

It's believed that food and energy prices spiked significantly during the 2021-2022 inflation surge. The difference between CPI and CPI-Core can help us investigate this claim, since CPI-Core excludes food and energy.

**Compute an array** containing the difference between CPI and CPI-Core for each month. This represents how much food and energy prices are contributing to overall inflation. (The first element of the array should correspond to the first row of `inflation`, and so on.)

*Note:* Use the original `inflation` table for this.

```python
food_energy_effect = ...
food_energy_effect
```

```python
grader.check("q1_4")
```

#### Question 1.5.

Add the `food_energy_effect` array as a column named `FoodEnergy` to the original `inflation` Table and sort the resulting Table by the `FoodEnergy` column in descending order. Call the new Table `by_food_energy`.

Try to do this with a single line of code, if you can.

```python
by_food_energy = ...
by_food_energy
```

```python
grader.check("q1_5")
```

#### Question 1.6.

It would be helpful to create a line plot of the food and energy effect over time. Unfortunately, the `Date` column is a string that can't be used as the independent variable of a line plot. You will need to create a new column, `Year`, that represents the year as a `float`. For example, January 1994 should be represented as `1994.0`, February 2000 should be represented as `2000.0833333333` (`0.0833333333` is $\\approx \\frac{1}{12}$), March 2020 would be represented as `2020.1666666667` and so on.

First, think about how you could quickly create an array of these `float` values using a `numpy` function. Remember, the first row of the data starts with January 1994 (`1994.0`) and each row increases the year value by one twelfth of a year. Store this array to the name `year_array`.

Then, add both the `year_array` array and the `food_energy_effect` array to the original `inflation` table. Label these columns `Year` and `FoodEnergy`.

Lastly, generate a line plot using one of the table methods you've learned in class.

```python
year_array = ...
inflation_with_year = ...
...
```

```python
grader.check("q1_6")
```

<!-- BEGIN QUESTION -->

#### Question 1.7.

Were food and energy prices contributing more to overall inflation during the 2021-2022 period compared to other times? Write an explanation of your opinion backed by the data displayed in the line plot created in the previous question, and any additional statistical analysis you perform. Cite specific historical events that align with key dates when possible to add context to your explanation.

<div class="alert alert-info" role="alert" style="color: black;">

Write your explanation here, replacing this text.

</div>

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Question 1.8.

A financial news reporter writes: "The CPI-Core measure is a better indicator of inflation because it removes the noise from volatile food and energy prices."

Do you agree or disagree with this statement? Consider: What are the advantages of using CPI-Core? What important information might be lost by excluding food and energy prices? For whom might food and energy price changes be especially important?

<div class="alert alert-info" role="alert" style="color: black;">

Write your response here, replacing this text.

</div>

<!-- END QUESTION -->

## 2. World Population

The following cell loads data about population changes in various countries between 2019 and 2020. The data includes births, deaths, and migration statistics.

```python
# Just run this cell
population = Table.read_table('world_population.csv')
population
```

#### Question 2.1.

Investigate the global birth *rate* during this time period. Assign `global_birth_rate` to the total birth rate across all countries in the dataset during this time interval. The annual birth rate for a year-long period is the total number of births in that period as a proportion of the population size at the start of the time period.

**Hint:** Which year corresponds to the start of the time period?

```python
total_births = ...
total_pop_2019 = ...
global_birth_rate = ...
global_birth_rate
```

```python
grader.check("q2_1")
```

#### Question 2.2.

Investigate the countries that had a relatively large percentage of their population move either into the country or out of the country. Use Table and array methods to determine the **number of countries** that had more than 0.5% of their population migrate in or out. The annual rate of migration for a year-long period is the net number of migrations (in minus out) as a proportion of the population size at the start of the period.

Start by creating an extension to the `population` Table named `migration_rates` that contains a new column with the value of the **annual *net* rate of migration** (as an absolute value) for each country. Then, determine how many countries had net migration rates higher than 0.5%.

```python
net_migration = ...
migration_rate = ...
migration_rates = ...
high_migration_countries = ...
high_migration_countries
```

```python
grader.check("q2_2")
```

#### Question 2.3.

Investigate the total number of births that occurred in European countries (Region 1) in 2020. Use Table and array methods to assign the total number of births, an integer, that occurred in region 1 (Europe) to `europe_births`.

```python
europe_births = ...
europe_births
```

```python
grader.check("q2_3")
```

#### Question 2.4.

Assign `less_than_europe_births` to the **number of countries** that had a total population in 2019 that was smaller than the **total *number of births* in region 1 (Europe)** during this time interval.

```python
less_than_europe_births = ...
less_than_europe_births
```

```python
grader.check("q2_4")
```

#### Question 2.5.

In the next question, you will be creating a visualization to understand the relationship between birth and death rates. The annual death rate for a year-long period is the total number of deaths in that period as a proportion of the population size at the start of the time period.

What visualization is most appropriate to see if there is an association between birth and death rates during a given time interval?

1. Scatter Plot
2. Line Graph
3. Vertical Bar Chart
4. Horizontal Bar Chart
5. Histogram

Assign `visualization` below to the number corresponding to the correct visualization.

```python
visualization = ...
```

```python
grader.check("q2_5")
```

#### Question 2.6.

In the code cell below, create a visualization that will help us determine if there is an association between birth rate and death rate during this time interval. You can take as many steps as you need, as you may find it helpful to create a new Table before creating the visualization.

```python
birth_rate = ...
death_rate = ...
rates = ...
...
```

<!-- BEGIN QUESTION -->

#### Question 2.7.

Is there an association between birth rate and death rate during this time interval? If so, how would you describe it? If not, how were you able to determine there was no association? **Write an explanation of your opinion backed by the data displayed in the visualization created in the previous question.** Offer an explanation as to why you believe it makes sense that there is or is not a correlation between these variables in context of the data.

<div class="alert alert-info" role="alert" style="color: black;">

Write your explanation here, replacing this text.

</div>

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Question 2.8.

The dataset you analyzed contains only 42 countries. Consider the following questions:

1. What limitations might arise from drawing global conclusions based on this sample of countries?
2. Looking at the countries included, do you notice any potential biases in which countries were selected?
3. How might your conclusions change if the dataset included all ~195 countries in the world?

<div class="alert alert-info" role="alert" style="color: black;">

Write your response here, replacing this text.

</div>

<!-- END QUESTION -->

## 3. Marginal Histograms

Consider the scatter plot below. It was created using a Table called `t` that has two columns: `x` and `y`.

<img src='scatter.png' width=500>

Below the scatter plot, there are two histograms. One is a histogram of the `x` column and one is a histogram of the `y` column.

| Histogram A | Histogram B |
| :---: | :---: |
|<img src='var1.png' width=350>|<img src='var2.png' width=350>|

#### Question 3.1.

Suppose we run `t.hist('x')`. Which histogram does this code produce? Assign `histogram_column_x` to either 1 or 2.

1. Histogram A
2. Histogram B

```python
histogram_column_x = ...
```

```python
grader.check("q3_1")
```

<!-- BEGIN QUESTION -->

#### Question 3.2.

State at least one reason why you chose the histogram from Question 3.1. Make sure to indicate which histogram you selected (ex: "I chose histogram A because ...") and cite one or more specific characteristics from the scatter plot to support your choice.

<div class="alert alert-info" role="alert" style="color: black;">

Write your explanation here, replacing this text.

</div>

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Question 3.3.

Imagine you were only shown Histogram A and Histogram B (without the scatter plot). Could you have predicted what the scatter plot would look like? What information is lost when we look at histograms of individual variables instead of the scatter plot showing both variables together?

<div class="alert alert-info" role="alert" style="color: black;">

Write your response here, replacing this text.

</div>

<!-- END QUESTION -->

## 4. Airbnb

Airbnb is a popular platform for short-term rentals. In this section, we'll investigate differences in stay duration between two major tourist cities: New York City and Paris.

The tables `nyc` and `paris` contain data on 200,000 Airbnb bookings in each city. Each row represents a single booking, and the `stay_nights` column records how many nights the guest stayed.

```python
# Just run this cell
nyc = Table.read_table('airbnb_nyc.csv')
paris = Table.read_table('airbnb_paris.csv')
print('NYC bookings:', nyc.num_rows)
print('Paris bookings:', paris.num_rows)
```

#### Question 4.1.

To start investigating the differences in these two cities, visualize the distribution of stay lengths in each of the cities.

Produce two histograms, one for each city, that shows the distribution of all stay lengths. Use the provided bin sizes that are stored to the array `equal_bins` and **assign the horizontal label to have the correct units**. Need to remember how? Reread [7.2.2: Histograms](https://inferentialthinking.com/chapters/07/2/Visualizing_Numerical_Distributions.html#histogram) in the textbook!

```python
equal_bins = np.arange(0, 32, 2)
```

```python
...
...
```

#### Question 4.2.

Assign `nyc_under_4` and `paris_under_4` to the percentage of stays that are less than 4 nights in their respective cities. The table below shows you the height of each bar in the histogram, measured in percent per night, for the first few bars of the histograms you produced in the previous question.

| bin | NYC | Paris |
|-----|-----|-------|
|[0, 2) | 3.5 | 2.1 |
|[2, 4) | 18.2 | 12.4 |
|[4, 6) | 15.8 | 13.9 |
|[6, 8) | 7.2 | 11.2 |

Your solution should be determined using only the information in this table, and you should not access the tables `nyc` and `paris` in any way.

```python
nyc_under_4 = ...
paris_under_4 = ...
print('NYC under 4 nights:', nyc_under_4, '%')
print('Paris under 4 nights:', paris_under_4, '%')
```

```python
grader.check("q4_2")
```

#### Question 4.3.

Let's take a closer look at the distribution of stay lengths in Paris. Assign `paris_median_bin` to an integer (1, 2, 3, or 4) that corresponds to the bin that contains the median stay length.

1. 0-4 nights
2. 4-8 nights
3. 8-12 nights
4. 12-16 nights

**Hint:** The median of a sorted list has half of the list elements to its left, and half to its right.

```python
paris_median_bin = ...
```

```python
grader.check("q4_3")
```

<!-- BEGIN QUESTION -->

#### Question 4.4.

Describe the main differences between the two histograms / the differences between the distributions of Airbnb stay lengths in these two cities. What about these two cities might be causing the differences in the distributions of stay lengths?

**Write an explanation** of your opinion backed by the data displayed in the histograms created in the previous questions, and any additional statistical analysis you perform. Cite specific qualities of these cities when possible to add context to your explanation.

<div class="alert alert-info" role="alert" style="color: black;">

Write your explanation here, replacing this text.

</div>

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Question 4.5.

A hotel industry analyst claims: "The histogram data proves that New Yorkers are less hospitable hosts than Parisians because guests leave sooner."

Evaluate this claim. What's wrong with this reasoning? What alternative explanations exist for the difference in stay lengths that have nothing to do with host hospitality?

<div class="alert alert-info" role="alert" style="color: black;">

Write your response here, replacing this text.

</div>

<!-- END QUESTION -->

# Submitting your work
You're done with this assignment! Assignments should be turned in using the following best practices:
1. Save your notebook.
2. Restart the kernel and run all cells up to this one.
3. Run the cell below with the code `grader.export(...)`. This will re-run all the tests. Make sure they are passing as you expect them to.
4. Download the file named `hw03_<date-time-stamp>.zip`, found in the explorer pane on the left side of the screen.
5. Upload `hw03_<date-time-stamp>.zip` to the corresponding assignment on Canvas.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(force_save=True, pdf=False)
```

Tests for q1_1 (points: None)
```
>>> isinstance(inflation, Table)
True
```
failure_message: ❌ inflation should be a Table
success_message: ✅ inflation is a Table


Tests for q1_2 (points: None)
```
>>> isinstance(by_cpi, Table) and isinstance(by_cpi_core, Table)
True
```
failure_message: ❌ by_cpi should be a Table
success_message: ✅ by_cpi is a Table


Tests for q1_3 (points: None)
```
>>> greatest_cpi.num_rows == 10
True
```
failure_message: ❌ greatest_cpi should have 10 rows
success_message: ✅ greatest_cpi has 10 rows


Tests for q1_4 (points: None)
```
>>> isinstance(food_energy_effect, np.ndarray)
True
```
failure_message: ❌ food_energy_effect should be an array
success_message: ✅ food_energy_effect is an array


Tests for q1_5 (points: None)
```
>>> 'FoodEnergy' in by_food_energy.labels
True
```
failure_message: ❌ by_food_energy should have a FoodEnergy column
success_message: ✅ by_food_energy has FoodEnergy column


Tests for q1_6 (points: None)
```
>>> np.isclose(year_array[0], 1994.0)
True
```
failure_message: ❌ year_array should start at 1994.0
success_message: ✅ year_array starts at 1994.0


Tests for q2_1 (points: None)
```
>>> 0 < global_birth_rate < 1
True
```
failure_message: ❌ global_birth_rate should be a proportion (between 0 and 1)
success_message: ✅ global_birth_rate is a number between 0 and 1


Tests for q2_2 (points: None)
```
>>> isinstance(high_migration_countries, (int, np.integer))
True
```
failure_message: ❌ high_migration_countries should be an integer
success_message: ✅ high_migration_countries is an integer


Tests for q2_3 (points: None)
```
>>> isinstance(europe_births, (int, np.integer, float, np.floating))
True
```
failure_message: ❌ europe_births should be a number
success_message: ✅ europe_births is a number


Tests for q2_4 (points: None)
```
>>> isinstance(less_than_europe_births, (int, np.integer))
True
```
failure_message: ❌ less_than_europe_births should be an integer
success_message: ✅ less_than_europe_births is an integer


Tests for q2_5 (points: None)
```
>>> visualization in [1, 2, 3, 4, 5]
True
```
failure_message: ❌ visualization should be 1, 2, 3, 4, or 5
success_message: ✅ You selected a valid choice


Tests for q3_1 (points: None)
```
>>> histogram_column_x in [1, 2]
True
```
failure_message: ❌ histogram_column_x should be 1 or 2
success_message: ✅ You selected a valid choice


Tests for q4_2 (points: None)
```
>>> isinstance(nyc_under_4, (int, float))
True
```
failure_message: ❌ nyc_under_4 should be a number
success_message: ✅ nyc_under_4 is a number


Tests for q4_3 (points: None)
```
>>> paris_median_bin in [1, 2, 3, 4]
True
```
failure_message: ❌ paris_median_bin should be 1, 2, 3, or 4
success_message: ✅ You selected a valid choice
