```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

# Homework 1: Causality and Expressions

Please complete this notebook by filling in the cells provided. Before you begin, be sure you've run the very first cell in this notebook to load `otter`, the auto-grading library that will provide feedback while completing this assignment. This cell should be run each time you load this notebook, as everything stored into memory is reset when you close Datahub for more than a few minutes.

**Recommended Reading:**
- [What is Data Science](http://www.inferentialthinking.com/chapters/01/what-is-data-science.html)
- [Causality and Experiments](http://www.inferentialthinking.com/chapters/02/causality-and-experiments.html) 
- [Programming in Python](http://www.inferentialthinking.com/chapters/03/programming-in-python.html)

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, please be sure to not re-assign variables throughout the notebook! For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

<div class="alert alert-warning" role="alert"  style="color: black;">
    
### Appropriate collaboration and AI usage

Consult with your teacher and course syllabus for information and policies regarding appropriate collaboration with other students, appropriate use of AI tools, and submission of late work.</div>

<div class="alert alert-danger" role="alert"  style="color: black;">
    
### This homework has hidden tests in it!

Additional tests will be run once your homework is submitted for grading. While you may pass all the tests you have access to before submission, you may not earn full credit if you do not pass the hidden tests as well.

Many of the tests you have access to before submitting only test to ensure you have given an answer that is formatted correctly and/or you have given an answer that *could* make sense in context. For example, a test you have access to while completing the assignment may check that you selected a valid choice for a multiple choice problem (1, 2, or 3) or that your answer is an integer between 0 and 50 if asked to count a subset of states in the United States. The tests that are run after submission will evaluate your work for accuracy. **Do not assume that just because all your tests pass before submission means that your answers are correct!**
</div>

## 1. Misleading Marketing

A fitness center advertisement claims:

> "Don't become a statistic! According to industry data, over 40% of gym members who cancel their memberships do so in the first five months of the year. Start your fitness journey in the summer instead!"

<!-- BEGIN QUESTION -->

#### Question 1.1.
This ad suggests that starting a gym membership in the first five months of the year (January through May) leads to a higher chance of cancellation compared to other times. Do the data in the ad support the claim that people who join gyms early in the year are more likely to cancel than those who join later? Please explain your answer.

**Note:** You can assume that "over 40%" means only slightly over. Had it been much higher, say closer to 50%, then the marketers would have said so.


<div class="alert alert-info" role="alert" style="color: black;">

Write your solution here, replacing this text.

</div>

<!-- END QUESTION -->

## 2. Analyzing Text in *Pride and Prejudice*

We can use Python to analyze the classic novel [*Pride and Prejudice*](https://www.gutenberg.org/ebooks/1342) by Jane Austen. The following code generates a scatter plot in which each dot corresponds to a chapter of the book. **Don't worry about how this code works for now! Just run it to see the scatter plot it produces**

```python
# Just run this cell
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')

# Read the text of Pride and Prejudice
with open('pride.txt', 'r') as file:
    pride_text = file.read()

# Split into chapters (the book uses "CHAPTER" as delimiter)
chapters = pride_text.split('CHAPTER ')[1:]  # Skip text before first chapter

# Count sentences (periods) and words in each chapter
sentences = [chapter.count('.') for chapter in chapters]
words = [len(chapter.split()) for chapter in chapters]

Table().with_columns(
    'Sentences',  sentences,
    'Words', words
    ).iscatter('Sentences')
```

The horizontal position of a dot measures the number of sentences in the chapter (approximated by counting periods). The vertical position measures the total number of words in the chapter. **You can use your mouse to hover over individual points and inspect their specific values**.

#### Question 2.1.
*Exactly* how many sentences are there in the chapter with the most words? Assign the answer as an integer to `text_q1`.

```python
text_q1 = ...
```

```python
grader.check("q2_1")
```

#### Question 2.2.
Which of the following chapters has the highest average number of words per sentence? Assign either 1, 2, or 3 to the name `text_q2` below.
1. The chapter with about 75 sentences
2. The chapter with about 125 sentences
3. The chapter with about 200 sentences


```python
text_q2 = ...
```

```python
grader.check("q2_2")
```

## 3. Names and Assignment Statements

#### Question 3.1.
When you run the following code, Python produces a cryptic error message:

```python
10 = 5 + 5
```

<div class="alert alert-danger" role="alert">
SyntaxError: cannot assign to literal here. Maybe you meant '==' instead of '='?
</div>

Choose the best explanation of what's wrong with the code, and then assign 1, 2, 3, or 4 to `error_q1` below to indicate your answer.

1. Python is not able to do addition.

2. The left-hand side of an assignment statement must be a variable name, not a number. Python cannot assign a value to `10`.

3. In Python, you cannot use `=` for assignment; you must use `==`.

4. The number 10 is already defined in Python, so you cannot redefine it.


```python
error_q1 = ...
```

```python
grader.check("q3_1")
```

#### Question 3.2.
When you run the following code, Python produces another cryptic error message:

```python
three = 3
nine = three times three
```

<div class="alert alert-danger" role="alert">

SyntaxError: invalid syntax

</div>

Choose the best explanation of what's wrong with the code, and then assign 1, 2, 3, or 4 to `error_q2` below to indicate your answer.

1. The word `times` is not a valid Python operator. To multiply, you must use `*`.

2. The variable name `nine` cannot be used because 9 is a number.

3. Python doesn't know how to multiply.

4. The variable `three` was never assigned a value, so it cannot be used.


```python
error_q2 = ...
```

```python
grader.check("q3_2")
```

#### Question 3.3.
When you run the following code, Python produces yet another cryptic error message:

```python
total cost = 50 + 25
```

<div class="alert alert-danger" role="alert">
SyntaxError: invalid syntax
</div>

Choose the best explanation of what's wrong with the code, and then assign 1, 2, or 3 to `error_q3` below to indicate your answer.

1. Python cannot add 50 and 25 because they are not stored in variables.

2. Variable names in Python cannot contain spaces. `total cost` should be written as `total_cost` or `totalcost`.

3. The word `total` is reserved in Python and cannot be used in variable names.


```python
error_q3 = ...
```

```python
grader.check("q3_3")
```

## 4. Text Message Vaccination Reminders

A [study](https://www.nature.com/articles/s41586-021-03843-2) conducted in collaboration with UCLA investigated whether simple text message reminders could increase vaccination rates. Researchers were interested in whether receiving reminder messages would cause more people to get vaccinated compared to receiving no reminders.

The study recruited approximately 90,000 pharmacy customers who had appointments scheduled for their first vaccination. Participants were randomly divided into two groups: about 45,000 received a series of text message reminders before their appointment, while the other 45,000 received no text messages (though they still had their appointments scheduled).

At the end of the study period, the researchers compared vaccination completion rates between the two groups.

#### Question 4.1.
Which statement best describes the *treatment* and *control* groups for this study? Assign either 1, 2, or 3 to the name `vaccine_q1` below.

1. The treatment group received text message reminders, and the control group received no text messages.

2. The treatment group was the 90,000 pharmacy customers recruited, and the control group was all other pharmacy customers not in the study.

3. There is no clear notion of *treatment* and *control* group in this study.


```python
vaccine_q1 = ...
```

```python
grader.check("q4_1")
```

#### Question 4.2.
Was this an observational study or a randomized controlled experiment? Assign either 1, 2, or 3 to the name `vaccine_q2` below.

1. This was an observational study.

2. This was a randomized controlled experiment.  

3. This was a randomized observational study.


```python
vaccine_q2 = ...
```

```python
grader.check("q4_2")
```

<!-- BEGIN QUESTION -->

#### Question 4.3.
The study found that text message reminders increased vaccination rates by about 2 percentage points. The researchers also measured whether people returned for their second dose and found similar improvements in the reminder group. Interestingly, the researchers tested different *types* of text messages (some emphasized ownership like "your vaccine is waiting for you," while others were simple reminders) and found that the ownership-focused messages worked slightly better.

Why do you think the researchers tested multiple types of text messages rather than just comparing "any text message" versus "no text message"? What additional insight does this provide? Please fully explain your answer.

<div class="alert alert-info" role="alert" style="color: black;">

Write your solution here, replacing this text.

</div>

<!-- END QUESTION -->

## 5. Changes in Streaming Subscriptions

A media analyst tracks streaming service subscriptions at a university. The table below displays the numbers of students subscribed to each streaming service in Fall 2019 and Fall 2024.

| Service          | Fall 2019    | Fall 2024   |
|------------------|--------------|-------------|
| Netflix          |      850     |    620      |
| Hulu             |      340     |    510      |
| Disney+          |      45      |    425      |


#### Question 5.1. 
Suppose you want to find the **biggest** absolute difference between the numbers of subscribers in the two years, among the three services.

In the cell below, compute this value and call it `biggest_change`. Use a single expression (a single line of code) to compute the answer. Let Python perform all the arithmetic (like subtracting 620 from 850) rather than simplifying the expression yourself. The built-in `abs` function takes a numerical input and returns the absolute value. The built-in `max` function can take in 3 arguments and returns the maximum of the three numbers.


```python
biggest_change = ...
biggest_change
```

```python
grader.check("q5_1")
```

#### Question 5.2.
Which of the three services had the **smallest** absolute difference? Assign `smallest_change_service` to 1, 2, or 3 where each number corresponds to the following service:

1: Netflix  
2: Hulu  
3: Disney+

Choose the number that corresponds to the service with the smallest absolute difference.

You should be able to answer by rough mental arithmetic, without having to calculate the exact value for each service. 


```python
smallest_change_service = ...
```

```python
grader.check("q5_2")
```

#### Question 5.3.
For each service, define the "relative change" to be the following: $\\large{\\frac{\\text{absolute difference}}{\\text{value in Fall 2019}} \\cdot 100}$ 

Fill in the code below such that `netflix_relative_change`, `hulu_relative_change` and `disney_relative_change` are assigned to the relative changes for their respective services.


```python
netflix_relative_change = ...
hulu_relative_change = ...
disney_relative_change = ...
print(f"Netflix: {netflix_relative_change:.2f}%")
print(f"Hulu: {hulu_relative_change:.2f}%")
print(f"Disney+: {disney_relative_change:.2f}%")
```

```python
grader.check("q5_3")
```

#### Question 5.4.
Assign `biggest_rel_change_service` to the service with the biggest relative change, where 1, 2, and 3 correspond to the following: 

1. Netflix 
2. Hulu  
3. Disney+

```python
biggest_rel_change_service = ...
```

```python
grader.check("q5_4")
```

## 6. Screen Time and Sleep Quality Study

A research team investigated the relationship between screen time before bed and sleep quality in teenagers. The study was motivated by concerns that blue light from devices might interfere with sleep.

<!-- BEGIN QUESTION -->

#### Question 6.1.
The data were gathered by the following procedure, reported in the study: "Between September and December 2023, teenagers aged 13-18 who visited a pediatric clinic for routine checkups completed a questionnaire about their typical screen use in the hour before bed and their sleep quality over the past month." 

Was this study observational, or was it a controlled experiment? Explain. 

<div class="alert alert-info" role="alert" style="color: black;">

Write your solution here, replacing this text.

</div>

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Question 6.2.
The study found that of the teenagers who used screens for more than 2 hours before bed, 68% reported poor sleep quality. Of those who used screens for 1-2 hours, 45% reported poor sleep quality. Of those who used screens for less than 1 hour before bed, 22% reported poor sleep quality. The study concluded that "Screen time in the hour before bed was strongly associated with poor sleep quality in teenagers."

Do the data support this statement? You may interpret "strongly" in any reasonable qualitative way.


<div class="alert alert-info" role="alert" style="color: black;">

Write your solution here, replacing this text.

</div>

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Question 6.3.
A news website reported on this study with the headline: "Scientists confirm: Phones before bed cause sleep problems in teens." Does the conclusion of the study claim that screen time causes poor sleep?


<div class="alert alert-info" role="alert" style="color: black;">

Write your solution here, replacing this text.

</div>

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

#### Question 6.4.
A sleep researcher pointed out that the study should have accounted for anxiety levels. Teenagers with higher anxiety are more likely to have trouble sleeping, and anxious teenagers might also use screens more as a coping mechanism or because they can't fall asleep anyway.

In what way does the knowledge of this possible link between anxiety and both screen time and sleep quality affect how we interpret the data from the study? Explain.


<div class="alert alert-info" role="alert" style="color: black;">

Write your solution here, replacing this text.

</div>

<!-- END QUESTION -->

# Submitting your work
You're done with this assignment! Assignments should be turned in using the following best practices:
1. Save your notebook.
2. Restart the kernel and run all cells up to this one.
3. Run the cell below with the code `grader.export(...)`. This will re-run all the tests. Make sure they are passing as you expect them to.
4. Download the file named `hw01_<date-time-stamp>.zip`, found in the explorer pane on the left side of the screen. **Note**: Clicking on the link in this notebook may result in an error, it's best to download from the file explorer panel.
5. Upload `hw01_<date-time-stamp>.zip` to the corresponding assignment on Canvas.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

```python
grader.export(force_save=True)
```

Tests for q2_1 (points: None)
```
>>> isinstance(text_q1, (int, float))
True
```
failure_message: ❌ Your approximation is not a number (integer or float)
success_message: ✅ You approximated using a number (integer or float)


Tests for q2_2 (points: None)
```
>>> text_q2 in [1, 2, 3]
True
```
failure_message: ❌ It looks like you didn't select 1, 2, or 3.
success_message: ✅ You selected a valid choice (1, 2, or 3)


Tests for q3_1 (points: None)
```
>>> error_q1 in [1, 2, 3, 4]
True
```
failure_message: ❌ It looks like you didn't select 1, 2, 3, or 4.
success_message: ✅ You selected a valid choice (1, 2, 3, or 4)


Tests for q3_2 (points: None)
```
>>> error_q2 in [1, 2, 3, 4]
True
```
failure_message: ❌ It looks like you didn't select 1, 2, 3, or 4.
success_message: ✅ You selected a valid choice (1, 2, 3, or 4)


Tests for q3_3 (points: None)
```
>>> error_q3 in [1, 2, 3]
True
```
failure_message: ❌ It looks like you didn't select 1, 2, or 3.
success_message: ✅ You selected a valid choice (1, 2, or 3)


Tests for q4_1 (points: None)
```
>>> vaccine_q1 in [1, 2, 3]
True
```
failure_message: ❌ It looks like you didn't select 1, 2, or 3.
success_message: ✅ You selected a valid choice (1, 2, or 3)


Tests for q4_2 (points: None)
```
>>> vaccine_q2 in [1, 2, 3]
True
```
failure_message: ❌ It looks like you didn't select 1, 2, or 3.
success_message: ✅ You selected a valid choice (1, 2, or 3)


Tests for q5_1 (points: None)
```
>>> isinstance(biggest_change, (int, float))
True
```
failure_message: ❌ biggest_change should be a number
success_message: ✅ biggest_change is a number


Tests for q5_2 (points: None)
```
>>> smallest_change_service in [1, 2, 3]
True
```
failure_message: ❌ It looks like you didn't select 1, 2, or 3.
success_message: ✅ You selected a valid choice (1, 2, or 3)


Tests for q5_3 (points: None)
```
>>> all((isinstance(x, (int, float)) for x in [netflix_relative_change, hulu_relative_change, disney_relative_change]))
True
```
failure_message: ❌ All relative changes should be numbers
success_message: ✅ All relative changes are numbers


Tests for q5_4 (points: None)
```
>>> biggest_rel_change_service in [1, 2, 3]
True
```
failure_message: ❌ It looks like you didn't select 1, 2, or 3.
success_message: ✅ You selected a valid choice (1, 2, or 3)
